# Un motor de ingesta sencillo para el dataset de vuelos

### Recuerda borrar todas las líneas que dicen `raise NotImplementedError`

El trabajo final consiste en la implementación guiada de un esqueleto de motor de ingesta para ficheros JSON recibidos diariamente con información de los vuelos que han tenido lugar en el día indicado en el nombre del fichero entre dos aeropuertos de los Estados Unidos.
* La estructura de cada JSON puede consultarse abriendo con un editor de texto cualquiera de los ficheros.
* El significado de cada una de las columnas puede consultarse en el fichero config.json que se encuentra en la carpeta config del repositorio.

Vamos a probar nuestro paquete del motor de ingesta. Antes de ejecutar este notebook, es imprescindible haber completado todo el código del motor de ingesta. 
El notebook solamente valida que el código del motor de ingesta con el cual se ha generado el paquete sea correcto. Para ello, el orden en el que debemos leer, entender y completar los ficheros del repositorio es:

1. Clase `MotorIngesta` en el fichero `motor_ingesta.py`. Sólo hay que completar la función `ingesta_fichero`
2. (Opcional) Completar los tests `test_aplana` y `test_ingesta_fichero` en el fichero `test_ingesta.py`
3. Clase `FlujoDiario` en el fichero `flujo_diario.py`. Completar primero el inicializador y luego el método `procesa_diario`. Al ir completando el código de este método, veremos que a su vez necesitamos completar las dos funciones siguientes.
4. Función `aniade_hora_utc` en el fichero `agregaciones.py`
5. (Opcional) El test de dicha función se llama  `test_aniade_hora_utc` en el fichero `test_ingesta.py`
6. Función `aniade_intervalos_por_aeropuerto` en el fichero `agregaciones.py`
7. (Opcional) El test de dicha función se llama `test_aniade_intervalos_por_aeropuerto` en el fichero `test_ingesta.py`

```python
!pip install loguru==0.7.1
```

Output:
```
Requirement already satisfied: loguru==0.7.1 in /mnt/c/Users/braya/OneDrive/Escritorio/spark_nticmaster/.venv-linux/lib/python3.10/site-packages (0.7.1)
```

## Instalamos el wheel en el cluster. 

```python
!pip install --force-reinstall ../dist/motor_ingesta-0.1.0-py3-none-any.whl
```

Output:
```
Processing /mnt/c/Users/braya/OneDrive/Escritorio/spark_nticmaster/dist/motor_ingesta-0.1.0-py3-none-any.whl
Installing collected packages: motor-ingesta
  Attempting uninstall: motor-ingesta
    Found existing installation: motor-ingesta 0.1.0
    Uninstalling motor-ingesta-0.1.0:
      Successfully uninstalled motor-ingesta-0.1.0
Successfully installed motor-ingesta-0.1.0
```

### Probamos la ingesta de un fichero

**Ejercicio 1 (2 puntos)**. Ingestar el fichero **`2023-01-01.json`** utilizando el motor de ingesta completo. Debe crearse el objeto de la clase MotorIngesta y utilizar el método `ingesta_fichero`, dejando el resultado en la variable `flights_df`. La variable `flujo_diario` contiene un objeto FlujoDiario inicializado con el path de configuración anterior, y lo vamos a usar ahora solo para leer adecuadamente la configuración y pasársela al objeto `motor_ingesta` como el argumento config.

* Este ejercicio requiere haber completado previamente el código del paquete de Python y haber generado el fichero .whl, y por tanto, la puntuación del ejercicio se debe a ese trabajo.

```python
# motor_ingesta/__init__.py
# flujo_diario.py — líneas 4 y 5
from motor_ingesta.motor_ingesta import MotorIngesta
from motor_ingesta.agregaciones import aniade_hora_utc, aniade_intervalos_por_aeropuerto
from motor_ingesta.flujo_diario import FlujoDiario

```

```python
# imports necesarios

path_config_flujo_diario = "../config/config.json"       # ruta del fichero config.json, que no pertenece al paquete
path_json_primer_dia = "../data/flights_json/landing/2023-01-01.json"           # ruta del fichero JSON de un día concreto que queremos ingestar, en nuestro caso 2023-01-01.json

flujo_diario = FlujoDiario(path_config_flujo_diario)
motor_ingesta = MotorIngesta(flujo_diario.config)
flights_df = motor_ingesta.ingesta_fichero(path_json_primer_dia)

# YOUR CODE HERE
# raise NotImplementedError
```

Output:
```
Enviroment: local
WARNING: Using incubator modules: jdk.incubator.vector
Using Spark's default log4j profile: org/apache/spark/log4j2-defaults.properties
26/04/12 21:29:50 WARN Utils: Your hostname, brrsanchezfi, resolves to a loopback address: 127.0.1.1; using 10.255.255.254 instead (on interface lo)
26/04/12 21:29:50 WARN Utils: Set SPARK_LOCAL_IP if you need to bind to another address
Using Spark's default log4j profile: org/apache/spark/log4j2-defaults.properties
Setting default log level to "WARN".
To adjust logging level use sc.setLogLevel(newLevel). For SparkR, use setLogLevel(newLevel).
26/04/12 21:29:51 WARN NativeCodeLoader: Unable to load native-hadoop library for your platform... using builtin-java classes where applicable
```

```python

flights_df.count()
```

Output:
```
15856
```

<!-- grade_id: ingesta-tests, points: 1, grade: true, solution: false -->
```python
assert(flights_df.count() == 15856)
# assert(len(flights_df.columns) == 18)
dtypes = dict(flights_df.dtypes)
assert(dtypes["Diverted"] == "boolean")
assert(dtypes["ArrTime"] == "int")
assert(flights_df.schema["Dest"].metadata == {"comment": "Destination Airport IATA code (3 letters)"})
```

<!-- grade_id: ingesta-tests, points: 1, grade: true, solution: false -->
```python
assert(flights_df.count() == 15856)
assert(len(flights_df.columns) == 18)
dtypes = dict(flights_df.dtypes)
assert(dtypes["Diverted"] == "boolean")
assert(dtypes["ArrTime"] == "int")
assert(flights_df.schema["Dest"].metadata == {"comment": "Destination Airport IATA code (3 letters)"})
```

### Probamos la función de añadir la hora en formato UTC

**Ejercicio 2 (2 puntos)** Probar la función de añadir hora UTC con el DF `flights_df` construido anteriormente. El resultado debe dejarse en la variable `flights_with_utc`. Recuerda que esto no es propiamente un test unitario.

* Este ejercicio requiere haber completado previamente el código de la función `aniade_hora_utc` del paquete de Python y haber generado el fichero .whl, y por tanto, la puntuación del ejercicio se debe a ese trabajo.

<!-- grade_id: flights-utc, points: 0, grade: false, solution: true -->
```python
# import necesarios

flights_with_utc = aniade_hora_utc(flujo_diario.spark,flights_df)

# YOUR CODE HERE
# raise NotImplementedError
```

```python
flights_with_utc.where("FlightTime is null").count()
```

Output:
```
264
```

```python
from pyspark.sql import functions as F
# assert(flights_with_utc.where("FlightTime is null").count() == 266)
types = dict(flights_with_utc.dtypes)
assert(flights_with_utc.dtypes[18] == ("FlightTime", "timestamp"))  # FlightTime debe ser la última columna

first_row = flights_with_utc.where("OriginAirportID = 12884").select(F.min("FlightTime").cast("string").alias("FlightTime")).first()
assert(first_row.FlightTime == "2023-01-01 10:59:00")
```

<!-- grade_id: flights-utc-tests, points: 2, grade: true, solution: false -->
```python
from pyspark.sql import functions as F
assert(flights_with_utc.where("FlightTime is null").count() == 266)
types = dict(flights_with_utc.dtypes)
assert(flights_with_utc.dtypes[18] == ("FlightTime", "timestamp"))  # FlightTime debe ser la última columna

first_row = flights_with_utc.where("OriginAirportID = 12884").select(F.min("FlightTime").cast("string").alias("FlightTime")).first()
assert(first_row.FlightTime == "2023-01-01 10:59:00")
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[65], line 2
      1 from pyspark.sql import functions as F
----> 2 assert(flights_with_utc.where("FlightTime is null").count() == 266)
      3 types = dict(flights_with_utc.dtypes)
      4 assert(flights_with_utc.dtypes[18] == ("FlightTime", "timestamp"))  # FlightTime debe ser la última columna

AssertionError:
```

## Probamos la función de añadir las columnas con la hora del siguiente vuelo, su aerolínea y el intervalo de tiempo transcurrido

**Ejercicio 3 (2.5 puntos)** Invocar a la función de añadir intervalos por aeropuerto, partiendo de la variable `flights_with_utc` del apartado anterior, dejando el resultado devuelto por la función en la variable `df_with_next_flight` cacheada.

* Este ejercicio requiere haber completado previamente el código de la función `aniade_intervalos_por_aeropuerto` en el paquete de Python y haber generado el fichero .whl, y por tanto, la puntuación del ejercicio se debe a ese trabajo.

<!-- grade_id: flights-intervalos, points: 0, grade: false, solution: true -->
```python
# imports necesarios

df_with_next_flight = aniade_intervalos_por_aeropuerto(flights_with_utc)

# YOUR CODE HERE
# raise NotImplementedError
```

<!-- grade_id: flights-intervalos-test, points: 2.5, grade: true, solution: false -->
```python
assert(df_with_next_flight.dtypes[19] == ("FlightTime_next", "timestamp"))
assert(df_with_next_flight.dtypes[20] == ("Airline_next", "string"))
assert(df_with_next_flight.dtypes[21] == ("diff_next", "bigint"))

first_row = df_with_next_flight.where("OriginAirportID = 12884")\
                               .select(F.col("FlightTime").cast("string"), 
                                       F.col("FlightTime_next").cast("string"), 
                                       F.col("Airline_next"),
                                       F.col("diff_next")).sort("FlightTime").first()

assert(first_row.FlightTime_next == "2023-01-01 16:36:00")
assert(first_row.Airline_next == "9E")
assert(first_row.diff_next == 20220)
```

## Corregimos que el último vuelo de cada aeropuerto y cada día tiene valor nulo en las 3 columnas `_next`

**Ejercicio 4 (2.5 puntos)**

Tal como está implementada la lógica del flujo diario, el último vuelo de cada día no tendrá informada la columna FlightTime_next porque no se dispone todavía de datos del día siguiente. Se pide **corregir este comportamiento** para solucionar los valores nulos, modificando el código del método `procesa_diario` de manera que, antes de escribir los datos del día actual, se hayan corregido las tres columnas `_next` en los datos del día anterior al que estamos ingestando. Una manera simple (aunque no necesariamente óptima) de conseguirlo es:
* Leer de la tabla la partición que se escribió el día previo, si existiera dicha tabla y dicha partición.
* Añadir al DF devuelto por `aniade_hora_utc` las 3 columnas que le faltan para tener la misma estructura que la tabla, que son `FlightTime_next`, `Airline_next` y `diff_next` (pueden ser en ese orden si la función `aniade_intervalos_por_aeropuerto` se ha implementado para añadirlas en ese orden), pero sin darles valor (con valor None, convirtiendo cada columna al tipo de dato adecuado para que después encaje con la tabla existente).
* Unir el DF del día previo y el que acabamos de calcular
* Invocar a `aniade_intervalos_por_aeropuerto` pasando como argumento el DF resultante de la unión.

Aparte de un test unitario (que se deja como optativo pero sin puntuación), la manera de comprobar el funcionamiento será invocar a `procesa_diario` del flujo diario, con los ficheros de dos días consecutivos, y después comprobar lo que se ha escrito en la tabla tras la ingesta del segundo fichero. Lo probaremos con los días 1 y 2 de enero de 2023.

* Este ejercicio requiere haber completado previamente el código del paquete de Python y haber generado el fichero .whl, y por tanto, la puntuación del ejercicio se debe a ese trabajo.

<!-- grade_id: procesa-diario, points: 0, grade: false, solution: true -->
```python
path_json_segundo_dia = "../data/flights_json/landing/2023-01-02.json"  # path del fichero 2023-01-02.json

flujo_diario.procesa_diario(path_json_primer_dia)
flujo_diario.procesa_diario(path_json_segundo_dia)

# Invoca al método procesa_diario del flujo con el path del fichero 2023-01-01.json
# Después invoca al método de nuevo con el path del fichero 2023-01-02.json

# YOUR CODE HERE
# raise NotImplementedError
```

Output:
```
26/04/12 21:30:29 WARN ObjectStore: Version information not found in metastore. hive.metastore.schema.verification is not enabled so recording the schema version 2.3.0
26/04/12 21:30:29 WARN ObjectStore: setMetaStoreSchemaVersion called but recording version is disabled: version = 2.3.0, comment = Set by MetaStore brrsanchezfi@127.0.1.1
26/04/12 21:30:30 WARN ObjectStore: Failed to get database default, returning NoSuchObjectException
[32m2026-04-12 21:30:33.363[0m | [1mINFO    [0m | [36mmotor_ingesta.flujo_diario[0m:[36mprocesa_diario[0m:[36m89[0m - [1mNo se han podido leer datos del día 2022-12-31: [TABLE_OR_VIEW_NOT_FOUND] The table or view `default`.`flights` cannot be found. Verify the spelling and correctness of the schema and catalog.
If you did not qualify the name with a schema, verify the current_schema() output, or qualify the name with the correct schema and catalog.
To tolerate the error on drop use DROP VIEW IF EXISTS or DROP TABLE IF EXISTS. SQLSTATE: 42P01;
'
```

```python
flujo_diario.spark.sql("select * from default.flights").show()
```

Output:
```
+-----------------+---------------+------+--------------+-----------+-------------+----+--------------------+---------+--------+-------+--------+-------+---------+--------+-------+--------+-------------------+-------------------+------------+---------+----------+
|Reporting_Airline|OriginAirportID|Origin|OriginCityName|OriginState|DestAirportID|Dest|        DestCityName|DestState|DepDelay|DepTime|ArrDelay|ArrTime|Cancelled|Diverted|AirTime|Distance|         FlightTime|    FlightTime_next|Airline_next|diff_next|FlightDate|
+-----------------+---------------+------+--------------+-----------+-------------+----+--------------------+---------+--------+-------+--------+-------+---------+--------+-------+--------+-------------------+-------------------+------------+---------+----------+
|               OO|          10155|   ACT|      Waco, TX|         TX|        11298| DFW|Dallas/Fort Worth...|       TX|      -2|    613|       0|    709|    false|   false|     31|      89|2023-01-02 12:13:00
```

<!-- grade_id: tests-unitarios, points: 2.5, grade: true, solution: false -->
```python
vuelos = spark.read.table("default.flights").sort("Origin", "FlightTime")
assert(vuelos.count() == 33931)
row = vuelos.where("FlightDate = '2023-01-01' and  Origin = 'ABE' and DepTime = 1734").first()
assert(row.diff_next == 44220)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[8], line 1
----> 1 vuelos = spark.read.table("default.flights").sort("Origin", "FlightTime")
      2 assert(vuelos.count() == 33931)
      3 row = vuelos.where("FlightDate = '2023-01-01' and  Origin = 'ABE' and DepTime = 1734").first()

NameError: name 'spark' is not defined
```

## Ejercicio opcional

**Ejercicio opcional (1 punto)** Completa los cuatro tests unitarios que encontrarás en el fichero `test_ingesta.py`. No tienes que escribir más código en este notebook.

- Se podrá optar a una calificación final de hasta 9.0 puntos sin resolver este ejercicio.